Прежде чем проверять задания:
- **Перезапустите ядро** (**restart the kernel**) (В меню, выбрать Ядро (Kernel) $\rightarrow$ Перезапустить (Restart)
- Затем **Выполнить все ячейки**  **run all cells** (В меню, выбрать Ячейка (Cell) $\rightarrow$ Запустить все (Run All).

Убедитесь, что заполнены все ячейки с комментарием "НАЧАЛО ВАШЕГО РЕШЕНИЯ".

После ячеек с заданием следуют ячейки с проверкой с помощью assert.

Если в коде есть ошибки, assert выведет уведомление об ошибке.

Если в коде нет ошибок, assert отработает без вывода дополнительной информации.

---

# Цель занятия
На этом занятии мы на практике закрепим работу с алгоритмами регрессии и классификации, изученными в рамках курса.
В задании необходимо имплементировать 8 моделей регрессии и 7 моделей классификации, а также подобрать к этим моделям оптимальные параметры с помощью GridSearchCV.

Scikit-learn (sklearn) и ряд других библиотек предлагают множество алгоритмов регрессии, которые можно использовать для решения задач прогнозирования численных значений. Вот некоторые из наиболее популярных алгоритмов регрессии:

1. Линейная регрессия (Linear Regression): `sklearn.linear_model.LinearRegression`
   - Простая модель, основанная на линейной комбинации признаков.
   - Подходит для задач с линейной зависимостью между признаками и целевой переменной.

2. Регрессия LASSO (LASSO Regression): `sklearn.linear_model.Lasso`
   - Линейная модель, которая также выполняет регуляризацию L1 (L1 regularization).
   - Помогает в выборе наиболее важных признаков и обеспечивает разреженность весов.

3. Регрессия Ridge (Ridge Regression): `sklearn.linear_model.Ridge`
   - Линейная модель, которая выполняет регуляризацию L2 (L2 regularization).
   - Помогает в уменьшении эффекта мультиколлинеарности и контролирует значения весов.

4. ElasticNet регрессия (ElasticNet Regression): `sklearn.linear_model.ElasticNet`
   - Комбинация регрессии LASSO и регрессии Ridge.
   - Обеспечивает баланс между отбором признаков и контролем мультиколлинеарности.

5. Регрессия по методу опорных векторов (Support Vector Regression): `sklearn.svm.SVR`
   - Расширение метода опорных векторов для задачи регрессии.
   - Позволяет работать с нелинейными зависимостями с помощью ядерных функций.

6. Решающие деревья (Decision Trees): `sklearn.tree.DecisionTreeRegressor`
   - Модель, основанная на деревьях принятия решений.
   - Позволяет обрабатывать нелинейные зависимости и автоматически выбирать наиболее важные признаки.

7. Случайный лес (Random Forest): `sklearn.ensemble.RandomForestRegressor`
   - Ансамблевая модель, состоящая из множества решающих деревьев.
   - Объединяет прогнозы множества деревьев для улучшения обобщающей способности.

8. Градиентный бустинг (Gradient Boosting): `sklearn.ensemble.GradientBoostingRegressor`
   - Ансамблевая модель, которая строит ансамбль слабых моделей в итеративном режиме.
   - Каждая новая модель исправляет ошибки предыдущих моделей, улучшая предсказания.

9. AdaBoost (Adaptive Boosting): `sklearn.ensemble.AdaBoostRegressor`
   - Алгоритм адаптивного бустинга, который взвешивает обучающие примеры для создания модели.
   - Фокусируется на примерах с большими ошибками для их исправления.

10. XGBoost: `xgboost.XGBRegressor`
    - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг.
    - Позволяет эффективно обрабатывать большие наборы данных и достигать высокой точности.

11. LightGBM: `lightgbm.LGBMRegressor`
    - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг.
    - Ориентирована на эффективность и скорость работы с большими объемами данных.

12. CatBoost: `catboost.CatBoostRegressor`
    - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг.
    - Обладает встроенной поддержкой категориальных признаков и автоматическим кодированием.
    
Некоторые из наиболее популярных алгоритмов классификации:

1. K-ближайших соседей (K-Nearest Neighbors): `sklearn.neighbors.KNeighborsClassifier`
   - Алгоритм, который классифицирует объекты на основе ближайших соседей в пространстве признаков.
   - Основывается на значении классов соседей для принятия решения о классификации.

2. Логистическая регрессия (Logistic Regression): `sklearn.linear_model.LogisticRegression`
   - Классификационная модель, которая использует логистическую функцию для предсказания вероятности принадлежности к классам.
   - Часто используется в задачах бинарной классификации.
   
3. Наивный Байес (Naive Bayes): `sklearn.naive_bayes.GaussianNB` (для непрерывных признаков) или `sklearn.naive_bayes.MultinomialNB` (для дискретных признаков)
   - Вероятностный классификатор, основанный на применении теоремы Байеса с предположением о независимости признаков.
   - Хорошо работает с наборами данных с большим числом признаков и предположением о независимости признаков.

4. Метод опорных векторов (Support Vector Machines): `sklearn.svm.SVC`
   - Алгоритм, который строит гиперплоскость или набор гиперплоскостей в многомерном пространстве для классификации объектов.
   - Может обрабатывать нелинейные зависимости с помощью ядерных функций.

5. Решающие деревья (Decision Trees): `sklearn.tree.DecisionTreeClassifier`
   - Модель, основанная на деревьях принятия решений.
   - Позволяет обрабатывать нелинейные зависимости и автоматически выбирать наиболее важные признаки.

6. Случайный лес (Random Forest): `sklearn.ensemble.RandomForestClassifier`
   - Ансамблевая модель, состоящая из множества решающих деревьев.
   - Объединяет прогнозы множества деревьев для улучшения обобщающей способности.

7. Градиентный бустинг (Gradient Boosting): `sklearn.ensemble.GradientBoostingClassifier`
   - Ансамблевая модель, которая строит ансамбль слабых моделей в итеративном режиме.
   - Каждая новая модель исправляет ошибки предыдущих моделей, улучшая предсказания.

8. AdaBoost (Adaptive Boosting): `sklearn.ensemble.AdaBoostClassifier`
   - Алгоритм адаптивного бустинга, который взвешивает обучающие примеры для создания модели.
   - Фокусируется на примерах с большими ошибками для их исправления.

9. XGBoost: `xgboost.XGBClassifier`
   - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг для классификации.
   - Позволяет эффективно обрабатывать большие наборы данных и достигать высокой точности.

10. LightGBM: `lightgbm.LGBMClassifier`
    - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг для классификации.
    - Ориентирована на эффективность и скорость работы с большими объемами данных.

11. CatBoost: `catboost.CatBoostClassifier`
    - Библиотека с открытым исходным кодом, предоставляющая градиентный бустинг для классификации.
    - Обладает встроенной поддержкой категориальных признаков и автоматическим кодированием.

Это некоторые из алгоритмов, доступных в scikit-learn и других библиотеках. Каждый алгоритм имеет свои особенности и подходит для разных типов данных и задач. Вам следует экспериментировать с разными алгоритмами и настраивать их параметры для достижения наилучших результатов в конкретной задаче.

При решении данного задания потребуются следующие библиотеки:
- pandas
- numpy
- sklearn
- xgboost
- catboost

Установить их можно следующим образом:
- pip install pandas==2.0.0
- pip install numpy==1.24.0
- pip install scikit-learn==1.2.0
- pip install xgboost==1.7.5
- pip install catboost==1.2

# Задача регрессии

Датасет "Boston" является одним из классических наборов данных. Он представляет собой набор данных, связанных с ценами на жилье в Бостоне, Массачусетс, США. Этот набор данных широко используется в задачах регрессии и демонстрирует взаимосвязь между различными факторами и стоимостью жилья.

Описание признаков (признаки ниже представлены в формате "название признака (тип данных)"):
1. CRIM (float): уровень преступности на душу населения по городам.
2. ZN (float): доля жилой земли, зонированной для участков более 25 000 кв.фут.
3. INDUS (float): доля нерозничных площадей на город.
4. CHAS (integer): Данный признак является фиктивной переменной, представляющей наличие или отсутствие ограничения участка жилья рекой Charles River.
5. NOX (float): концентрация оксидов азота (частей на 10 миллионов).
6. RM (float): среднее количество комнат в жилом помещении.
7. AGE (float): доля занятых владельцами единиц, построенных до 1940 года.
8. DIS (float): взвешенные расстояния до пяти бостонских центров занятости.
9. RAD (integer): индекс доступности к радиальным магистралям.
10. TAX (integer): полная ставка налога на имущество на 10 000 долларов США.
11. PTRATIO (float): соотношение учеников и учителей по городам.
12. B (float): 1000 * (Bk - 0.63) ^ 2, где Bk - доля чернокожего населения по городам.
13. LSTAT (float): процент более низкого статуса населения.
14. MEDV (float): медианная стоимость домов, занимаемых владельцами, в тысячах долларов.

Целевая переменная (выходной признак):
- MEDV: медианная стоимость домов, занимаемых владельцами, в тысячах долларов.

Цель состоит в том, чтобы предсказать медианную стоимость жилья на основе доступных признаков. Этот набор данных часто используется для задачи регрессии и оценки моделей машинного обучения на основе различных признаков, связанных с жильем.

Мы будем использовать 8 моделей регрессии и GridSearchCV для подбора параметров:
1. Линейная регрессия (Linear Regression): `sklearn.linear_model.LinearRegression`
2. Регрессия по методу опорных векторов (Support Vector Regression): `sklearn.svm.SVR`
3. Решающие деревья (Decision Trees): `sklearn.tree.DecisionTreeRegressor`
4. Случайный лес (Random Forest): `sklearn.ensemble.RandomForestRegressor`
5. Градиентный бустинг (Gradient Boosting): `sklearn.ensemble.GradientBoostingRegressor`
6. AdaBoost (Adaptive Boosting): `sklearn.ensemble.AdaBoostRegressor`
7. XGBoost: `xgboost.XGBRegressor`
8. CatBoost: `catboost.CatBoostRegressor`

```python
import pandas as pd
import numpy as np
import time

from sklearn.model_selection import train_test_split
from sklearn.compose import ColumnTransformer
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler, OneHotEncoder

from sklearn.model_selection import GridSearchCV
from sklearn.metrics import mean_squared_error
from sklearn.metrics import f1_score

import xgboost as xgb
import catboost as cb

import warnings
warnings.filterwarnings("ignore")
```

```python
# Загрузка данных из CSV файла
data = pd.read_csv("boston.csv", delimiter=",")
data.head()
```

Output:
```
CRIM    ZN  INDUS  CHAS     NX     RM   AGE     DIS  RAD    TAX  \
0  0.00632  18.0   2.31     0  0.538  6.575  65.2  4.0900    1  296.0   
1  0.02731   0.0   7.07     0  0.469  6.421  78.9  4.9671    2  242.0   
2  0.02729   0.0   7.07     0  0.469  7.185  61.1  4.9671    2  242.0   
3  0.03237   0.0   2.18     0  0.458  6.998  45.8  6.0622    3  222.0   
4  0.06905   0.0   2.18     0  0.458  7.147  54.2  6.0622    3  222.0   

   PTRATIO       B  LSTAT  MEDV  
0     15.3  396.90   4.98  24.0  
1     17.8  396.90   9.14  21.6  
2     17.8  392.83   4.03  34.7  
3     18.7  394.63   2.94  33.4  
4     18.7  396.90   5.33  36.2
```

```python
# Разделение на признаки (X) и целевую переменную (y)
X = data.drop('MEDV', axis=1)
y = data['MEDV']


# Разделение на обучающий и тестовый наборы данных
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

# Определение препроцессора для категориальных и числовых признаков
categorical_features = ['RAD']
categorical_transformer = OneHotEncoder(handle_unknown='ignore')
numerical_features = X_train.select_dtypes(include=['int64', 'float64']).columns
numerical_transformer = StandardScaler()

# Объединение препроцессоров
preprocessor = ColumnTransformer(
    transformers=[
        ('cat', categorical_transformer, categorical_features),
        ('num', numerical_transformer, numerical_features)
    ])

# Создание пайплайна для препроцессора
preprocessing_pipeline = Pipeline(steps=[('preprocessor', preprocessor)])

# Применение препроцессора к данным
X_train = preprocessing_pipeline.fit_transform(X_train)
X_test = preprocessing_pipeline.transform(X_test)
```

<!-- grade_id: cell-a3a0470231ab2bc8, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
#Задание 1.1 (1 балл)
model = LinearRegression()

# Определение параметров для GridSearchCV
params = {'fit_intercept': [True, False]}

# Создание объекта GridSearchCV
#Задание 1.2 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров LinearRegression
LinearRegression_best_params = grid_search.best_params_
print("LinearRegression_best_params: ", LinearRegression_best_params)

# Вычисление MSE
LinearRegression_mse = mean_squared_error(y_test, y_pred)
print("LinearRegression_mse: ", LinearRegression_mse)
```

Output:
```
LinearRegression_best_params:  {'fit_intercept': True}
LinearRegression_mse:  24.81844273848132
```

<!-- grade_id: cell-3b4da2f7c279a8d3, points: 2, grade: true, solution: false -->
```python
assert type(model) == LinearRegression
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-ee967bd36c165b8b, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVR

#Задание 1.3 (1 балл)
model = SVR()

# Определение параметров для GridSearchCV
params = {'C': [0.1, 1, 10], 'kernel': ['linear', 'rbf', 'poly'], 'gamma': ['scale', 'auto']}

# Создание объекта GridSearchCV
#Задание 1.4 (1 балл)
grid_search = GridSearchCV(estimator=model, param_grid=params,scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров LinearRegression
SVR_best_params = grid_search.best_params_
print("SVR_best_params: ", SVR_best_params)

# Вычисление MSE
SVR_mse = mean_squared_error(y_test, y_pred)
print("SVR_mse: ", SVR_mse)
```

Output:
```
SVR_best_params:  {'C': 10, 'gamma': 'scale', 'kernel': 'rbf'}
SVR_mse:  13.497844888204206
```

<!-- grade_id: cell-d8eabe616f47c022, points: 2, grade: true, solution: false -->
```python
assert type(model) == SVR
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-e56216bd2ec299b5, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
#Задание 1.5 (1 балл)
model = DecisionTreeRegressor ()

# Определение параметров для GridSearchCV

params = {'max_depth': [None, 5, 10]}

# Создание объекта GridSearchCV
#Задание 1.6 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров LinearRegression
DecisionTreeRegressor_best_params = grid_search.best_params_
print("DecisionTreeRegressor_best_params: ", DecisionTreeRegressor_best_params)

# Вычисление MSE
DecisionTreeRegressor_mse = mean_squared_error(y_test, y_pred)
print("DecisionTreeRegressor_mse: ", DecisionTreeRegressor_mse)
```

Output:
```
DecisionTreeRegressor_best_params:  {'max_depth': None}
DecisionTreeRegressor_mse:  21.316274509803918
```

<!-- grade_id: cell-e76219089ca29b9f, points: 2, grade: true, solution: false -->
```python
assert type(model) == DecisionTreeRegressor
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-2de3053b178599de, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

# Создание объекта RandomForestRegressor
#Задание 1.7 (1 балл)
model = RandomForestRegressor()

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [100, 200, 300],
    'max_depth': [None, 5, 10],
    'min_samples_split': [2, 5, 10]
}

# Создание объекта GridSearchCV
#Задание 1.8 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров RandomForestRegressor
RandomForestRegressor_best_params = grid_search.best_params_
print("RandomForestRegressor_best_params: ", RandomForestRegressor_best_params)

# Вычисление MSE
RandomForestRegressor_mse = mean_squared_error(y_test, y_pred)
print("RandomForestRegressor_mse: ", RandomForestRegressor_mse)
```

Output:
```
RandomForestRegressor_best_params:  {'max_depth': 10, 'min_samples_split': 2, 'n_estimators': 200}
RandomForestRegressor_mse:  8.373546319204586
```

<!-- grade_id: cell-243b0bf409d92646, points: 2, grade: true, solution: false -->
```python
assert type(model) == RandomForestRegressor
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-5ce5ce26909883fd, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import GradientBoostingRegressor

# Создание объекта GradientBoostingRegressor
#Задание 1.9 (1 балл)
model = GradientBoostingRegressor()

# Определение параметров для GridSearchCV

params = {
    'n_estimators': [100, 200, 300],
    'learning_rate': [0.1, 0.01, 0.001],
    'max_depth': [3, 5, 10]
}

# Создание объекта GridSearchCV
#Задание 1.10 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров GradientBoostingRegressor
GradientBoostingRegressor_best_params = grid_search.best_params_
print("GradientBoostingRegressor_best_params: ", GradientBoostingRegressor_best_params)

# Вычисление MSE
GradientBoostingRegressor_mse = mean_squared_error(y_test, y_pred)
print("GradientBoostingRegressor_mse: ", GradientBoostingRegressor_mse)
```

Output:
```
GradientBoostingRegressor_best_params:  {'learning_rate': 0.1, 'max_depth': 3, 'n_estimators': 100}
GradientBoostingRegressor_mse:  6.470048449756426
```

<!-- grade_id: cell-c20b00f45cf2781c, points: 2, grade: true, solution: false -->
```python
assert type(model) == GradientBoostingRegressor
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-65d703159cd1eda3, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import AdaBoostRegressor

# Создание объекта AdaBoostRegressor
#Задание 1.11 (1 балл)
model = AdaBoostRegressor()

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [50, 100, 200],
    'learning_rate': [0.1, 0.01, 0.001]
}

# Создание объекта GridSearchCV
#Задание 1.12 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров AdaBoostRegressor
AdaBoostRegressor_best_params = grid_search.best_params_
print("AdaBoostRegressor_best_params: ", AdaBoostRegressor_best_params)

# Вычисление MSE
AdaBoostRegressor_mse = mean_squared_error(y_test, y_pred)
print("AdaBoostRegressor_mse: ", AdaBoostRegressor_mse)
```

Output:
```
AdaBoostRegressor_best_params:  {'learning_rate': 0.1, 'n_estimators': 200}
AdaBoostRegressor_mse:  12.856115717413791
```

<!-- grade_id: cell-603495670d24c3c7, points: 2, grade: true, solution: false -->
```python
assert type(model) == AdaBoostRegressor
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-61b471f7697689a9, points: 0, grade: false, solution: true -->
```python
from xgboost import XGBRegressor

# Создание объекта XGBRegressor
#Задание 1.13 (1 балл)
model = XGBRegressor ()

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [100, 200, 300],
    'learning_rate': [0.1, 0.01, 0.001],
    'max_depth': [3, 5, 10]
}

# Создание объекта GridSearchCV
#Задание 1.14 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров XGBRegressor
XGBRegressor_best_params = grid_search.best_params_
print("XGBRegressor_best_params: ", XGBRegressor_best_params)

# Вычисление MSE
XGBRegressor_mse = mean_squared_error(y_test, y_pred)
print("XGBRegressor_mse: ", XGBRegressor_mse)
```

Output:
```
XGBRegressor_best_params:  {'learning_rate': 0.1, 'max_depth': 3, 'n_estimators': 200}
XGBRegressor_mse:  5.988378570059145
```

<!-- grade_id: cell-7eb8f4527128f36b, points: 2, grade: true, solution: false -->
```python
assert type(model) == XGBRegressor
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-0382c850da183c6e, points: 0, grade: false, solution: true -->
```python
from catboost import CatBoostRegressor

# Создание объекта CatBoostRegressor
#Задание 1.15 (1 балл)
model = CatBoostRegressor ()

# Определение параметров для GridSearchCV
params = {
    'learning_rate': [0.1, 0.01, 0.001],
    'n_estimators': [50, 100, 200],
    'max_depth': [3, 5, 10],
    'subsample': [1.0],
    'colsample_bylevel': [1.0],
    'reg_lambda': [3.0]
}

# Создание объекта GridSearchCV
#Задание 1.16 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='neg_mean_squared_error')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров CatBoostRegressor
CatBoostRegressor_best_params = grid_search.best_params_
print("CatBoostRegressor_best_params: ", CatBoostRegressor_best_params)

# Вычисление MSE
CatBoostRegressor_mse = mean_squared_error(y_test, y_pred)
print("CatBoostRegressor_mse: ", CatBoostRegressor_mse)
```

Output:
```
0:	learn: 8.8884998	total: 817us	remaining: 40ms
1:	learn: 8.4867901	total: 1.9ms	remaining: 45.7ms
2:	learn: 8.1159104	total: 4.01ms	remaining: 62.9ms
3:	learn: 7.7329591	total: 5.18ms	remaining: 59.6ms
4:	learn: 7.3460791	total: 6.96ms	remaining: 62.6ms
5:	learn: 7.0273671	total: 7.92ms	remaining: 58.1ms
6:	learn: 6.7302380	total: 9.39ms	remaining: 57.7ms
7:	learn: 6.3963844	total: 10.4ms	remaining: 54.6ms
8:	learn: 6.1013672	total: 11ms	remaining: 50ms
9:	learn: 5.8595928	total: 11.7ms	remaining: 46.8ms
10:	learn: 5.6345369	total: 12.7ms	remaining: 45.2ms
11:	learn: 5.3942623	total: 13.6ms	remaining: 43.2ms
12:	learn: 5.2391727	total: 15.1ms	remaining: 43ms
13:	learn: 5.1001731	total: 15.7ms	remaining: 40.3ms
14:	learn: 4.9052527	total: 16.3ms	remaining: 38.1ms
15:	learn: 4.7422603	total: 17.1ms	remaining: 36.2ms
16:	learn: 4.5967415	total: 17.6ms	remaining: 34.2ms
17:	learn: 4.4671993	total: 18.2ms	remaining: 32.3ms
18:	learn: 4.3827335	total: 18.7ms	remaining: 30.6ms
19:	learn: 4.
```

<!-- grade_id: cell-3e124666cb4b7f82, points: 2, grade: true, solution: false -->
```python
assert type(model) == CatBoostRegressor
assert type(grid_search) == GridSearchCV
```

```python
# Вывод результатов моделей регрессии 

data = [ 
    ['LinearRegression', f'{LinearRegression_mse}', f'{LinearRegression_best_params}'],
    ['SVR', f'{SVR_mse}', f'{SVR_best_params}'],
    ['DecisionTreeRegressor', f'{DecisionTreeRegressor_mse}', f'{DecisionTreeRegressor_best_params}'],
    ['RandomForestRegressor', f'{RandomForestRegressor_mse}', f'{RandomForestRegressor_best_params}'],
    ['GradientBoostingRegressor', f'{GradientBoostingRegressor_mse}', f'{GradientBoostingRegressor_best_params}'],
    ['AdaBoostRegressor', f'{AdaBoostRegressor_mse}', f'{AdaBoostRegressor_best_params}'],
    ['XGBRegressor', f'{XGBRegressor_mse}', f'{XGBRegressor_best_params}'],
    ['CatBoostRegressor', f'{CatBoostRegressor_mse}', f'{CatBoostRegressor_best_params}']
]


# Создание DataFrame из данных
df = pd.DataFrame(data, columns=['Модель', 'MSE', 'Параметры'], index=range(1, 1 + len(data)))
pd.set_option('display.max_colwidth', None)
# Отображение DataFrame
df.head(len(data))
```

Output:
```
Модель                 MSE  \
1           LinearRegression   24.81844273848132   
2                        SVR  13.497844888204206   
3      DecisionTreeRegressor  21.316274509803918   
4      RandomForestRegressor   8.373546319204586   
5  GradientBoostingRegressor   6.470048449756426   
6          AdaBoostRegressor  12.856115717413791   
7               XGBRegressor   5.988378570059145   
8          CatBoostRegressor  10.905647311507755   

                                                                                                                    Параметры  
1                                                                                                     {'fit_intercept': True}  
2                                                                                {'C': 10, 'gamma': 'scale', 'kernel': 'rbf'}  
3                                                                                                         {'max_depth': None}  
4                                        
```

# Задача классификации

Мы будем использовать датасет "Bank Marketing". В датасете "Bank Marketing" содержатся данные, связанные с маркетинговыми кампаниями банка. Вот колонки, представленные в датасете "Bank Marketing":

1. **age**: Возраст клиента.

2. **job**: Тип работы клиента.

3. **marital**: Семейное положение клиента.

4. **education**: Образование клиента.

5. **default**: Есть ли у клиента кредитный дефолт (задолженность).

6. **balance**: Баланс на счете клиента.

7. **housing**: Имеет ли клиент ипотеку (жилую недвижимость).

8. **loan**: Имеет ли клиент личный заем.

9. **contact**: Способ связи с клиентом.

10. **day**: День месяца, когда был выполнен последний контакт.

11. **month**: Месяц года, когда был выполнен последний контакт.

12. **duration**: Продолжительность последнего контакта в секундах.

13. **campaign**: Количество контактов, выполненных во время данной кампании.

14. **pdays**: Количество дней, прошедших с момента последнего контакта до предыдущей кампании.

15. **previous**: Количество контактов, выполненных перед текущей кампанией.

16. **poutcome**: Результат предыдущей маркетинговой кампании.

17. **y**: Целевая переменная, которую необходимо предсказать. Например, это может быть информация о том, стал ли клиент подписчиком (1) или нет (0) на определенное предложение банка.

Мы будем использовать 7 моделей классификации и GridSearchCV для подбора параметров:
1. Логистическая регрессия (Logistic Regression): `sklearn.linear_model.LogisticRegression`
2. Метод опорных векторов (Support Vector Machines): `sklearn.svm.SVC`
3. Решающие деревья (Decision Trees): `sklearn.tree.DecisionTreeClassifier`
4. Случайный лес (Random Forest): `sklearn.ensemble.RandomForestClassifier`
5. Градиентный бустинг (Gradient Boosting): `sklearn.ensemble.GradientBoostingClassifier`
6. XGBoost: `xgboost.XGBClassifier`
7. CatBoost: `catboost.CatBoostClassifier`

```python
# Загрузка данных из CSV файла
data = pd.read_csv("bank-full.csv", delimiter=";")
data.head()
```

Output:
```
age           job  marital  education default  balance housing loan  \
0   58    management  married   tertiary      no     2143     yes   no   
1   44    technician   single  secondary      no       29     yes   no   
2   33  entrepreneur  married  secondary      no        2     yes  yes   
3   47   blue-collar  married    unknown      no     1506     yes   no   
4   33       unknown   single    unknown      no        1      no   no   

   contact  day month  duration  campaign  pdays  previous poutcome   y  
0  unknown    5   may       261         1     -1         0  unknown  no  
1  unknown    5   may       151         1     -1         0  unknown  no  
2  unknown    5   may        76         1     -1         0  unknown  no  
3  unknown    5   may        92         1     -1         0  unknown  no  
4  unknown    5   may       198         1     -1         0  unknown  no
```

```python
# Разделение на признаки (X) и целевую переменную (y)
X = data.drop('y', axis=1)
y = data['y'].apply(lambda x: 1 if x=="yes" else 0)


# Разделение на обучающий и тестовый наборы данных
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

# Определение препроцессора для категориальных и числовых признаков
categorical_features = ['job', 'marital', 'education', 'default', 'housing', 'loan', 'contact', 'month', 'poutcome']
categorical_transformer = OneHotEncoder(handle_unknown='ignore')
numerical_features = X_train.select_dtypes(include=['int64', 'float64']).columns
numerical_transformer = StandardScaler()

# Объединение препроцессоров
preprocessor = ColumnTransformer(
    transformers=[
        ('cat', categorical_transformer, categorical_features),
        ('num', numerical_transformer, numerical_features)
    ])

# Создание пайплайна для препроцессора
preprocessing_pipeline = Pipeline(steps=[('preprocessor', preprocessor)])

# Применение препроцессора к данным
X_train = preprocessing_pipeline.fit_transform(X_train)
X_test = preprocessing_pipeline.transform(X_test)
```

<!-- grade_id: cell-b36ae883d67c44b4, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LogisticRegression

#Задание 2.1 (1 балл)
model = LogisticRegression()

# Определение параметров для GridSearchCV
params = {
'penalty': ['l1', 'l2'],
'C': [0.1, 1.0, 10.0]
}

# Создание объекта GridSearchCV
#Задание 2.2 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров LogisticRegression
LogisticRegression_best_params = grid_search.best_params_
print("LogisticRegression_best_params: ", LogisticRegression_best_params)

# Вычисление F1-меры
LogisticRegression_f1 = f1_score(y_test, y_pred, average='macro')
print("LogisticRegression_f1: ", LogisticRegression_f1)
```

Output:
```
LogisticRegression_best_params:  {'C': 10.0, 'penalty': 'l2'}
LogisticRegression_f1:  0.6962137604356302
```

<!-- grade_id: cell-e8f682ad6574994e, points: 2, grade: true, solution: false -->
```python
assert type(model) == LogisticRegression
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-07c7a82699aba882, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

#Задание 2.3 (1 балл)
model = SVC()

# Определение параметров для GridSearchCV
params = {
    'kernel': ['linear'],
    'C': [1.0]
}

# Создание объекта GridSearchCV
#Задание 2.4 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров SVC
SVC_best_params = grid_search.best_params_
print("SVC_best_params: ", SVC_best_params)

# Вычисление F1-меры
SVC_f1 = f1_score(y_test, y_pred, average='macro')
print("SVC_f1: ", SVC_f1)
```

Output:
```
SVC_best_params:  {'C': 1.0, 'kernel': 'linear'}
SVC_f1:  0.615523472539518
```

<!-- grade_id: cell-e7f61e84459f0c86, points: 2, grade: true, solution: false -->
```python
assert type(model) == SVC
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-9028c34e4703f7f8, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeClassifier

#Задание 2.5 (1 балл)
model = DecisionTreeClassifier()

# Определение параметров для GridSearchCV
params = {
    'criterion': ['gini', 'entropy'],
    'max_depth': [None, 5, 10]
}

# Создание объекта GridSearchCV
#Задание 2.6 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров DecisionTreeClassifier
DecisionTreeClassifier_best_params = grid_search.best_params_
print("DecisionTreeClassifier_best_params: ", DecisionTreeClassifier_best_params)

# Вычисление F1-меры
DecisionTreeClassifier_f1 = f1_score(y_test, y_pred, average='macro')
print("DecisionTreeClassifier_f1: ", DecisionTreeClassifier_f1)
```

Output:
```
DecisionTreeClassifier_best_params:  {'criterion': 'gini', 'max_depth': 10}
DecisionTreeClassifier_f1:  0.7150713736547856
```

<!-- grade_id: cell-e5f4270f3ee48fa8, points: 2, grade: true, solution: false -->
```python
assert type(model) == DecisionTreeClassifier
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-3a383c70d0b572c3, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
#Задание 2.7 (1 балл)
model = RandomForestClassifier()

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [100, 200, 300],
    'max_depth': [None, 5, 10]
}

# Создание объекта GridSearchCV
#Задание 2.8 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров RandomForestClassifier
RandomForestClassifier_best_params = grid_search.best_params_
print("RandomForestClassifier_best_params: ", RandomForestClassifier_best_params)

# Вычисление F1-меры
RandomForestClassifier_f1 = f1_score(y_test, y_pred, average='macro')
print("RandomForestClassifier_f1: ", RandomForestClassifier_f1)
```

Output:
```
RandomForestClassifier_best_params:  {'max_depth': None, 'n_estimators': 200}
RandomForestClassifier_f1:  0.7290506784614048
```

<!-- grade_id: cell-1558a47cc9230343, points: 2, grade: true, solution: false -->
```python
assert type(model) == RandomForestClassifier
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-87c99b53ef046ce5, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import GradientBoostingClassifier

#Задание 2.9 (1 балл)
model = GradientBoostingClassifier()

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [30, 50, 100],
    'max_depth': [5],
    'learning_rate': [0.1, 0.01]
}

# Создание объекта GridSearchCV
#Задание 2.10 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров GradientBoostingClassifier
GradientBoostingClassifier_best_params = grid_search.best_params_
print("GradientBoostingClassifier_best_params: ", GradientBoostingClassifier_best_params)

# Вычисление F1-меры
GradientBoostingClassifier_f1 = f1_score(y_test, y_pred, average='macro')
print("GradientBoostingClassifier_f1: ", GradientBoostingClassifier_f1)
```

Output:
```
GradientBoostingClassifier_best_params:  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}
GradientBoostingClassifier_f1:  0.7511473381048341
```

<!-- grade_id: cell-eef10fd55b17650e, points: 2, grade: true, solution: false -->
```python
assert type(model) == GradientBoostingClassifier
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-9ff51eaa71d964a9, points: 0, grade: false, solution: true -->
```python
import xgboost as xgb

#Задание 2.11 (1 балл)
model = xgb.XGBClassifier ()

# Определение параметров для GridSearchCV

params = {
    'n_estimators': [30, 50, 100],
    'max_depth': [3, 5],
    'learning_rate': [0.1, 0.01]
}

# Создание объекта GridSearchCV
#Задание 2.12 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров XGBClassifier
XGBClassifier_best_params = grid_search.best_params_
print("XGBClassifier_best_params: ", XGBClassifier_best_params)

# Вычисление F1-меры
XGBClassifier_f1 = f1_score(y_test, y_pred, average='macro')
print("XGBClassifier_f1: ", XGBClassifier_f1)
```

Output:
```
XGBClassifier_best_params:  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}
XGBClassifier_f1:  0.7482487047186519
```

<!-- grade_id: cell-265ff171f6dfff58, points: 2, grade: true, solution: false -->
```python
assert type(model) == xgb.XGBClassifier
assert type(grid_search) == GridSearchCV
```

<!-- grade_id: cell-e49e7d38996f3ac1, points: 0, grade: false, solution: true -->
```python
import catboost as cb

#Задание 2.13 (1 балл)
model = cb.CatBoostClassifier(silent=True)

# Определение параметров для GridSearchCV
params = {
    'n_estimators': [30, 50, 100],
    'max_depth': [3, 5],
    'learning_rate': [0.1, 0.01]
}

# Создание объекта GridSearchCV
#Задание 2.14 (1 балл)
grid_search = GridSearchCV (estimator=model, param_grid=params, scoring='f1_macro')

# Обучение модели на обучающих данных
grid_search.fit(X_train, y_train)

# Получение лучшей модели
best_model = grid_search.best_estimator_

# Предсказание на тестовых данных
y_pred = best_model.predict(X_test)

# Получение лучших параметров CatBoostClassifier
CatBoostClassifier_best_params = grid_search.best_params_
print("CatBoostClassifier_best_params: ", CatBoostClassifier_best_params)

# Вычисление F1-меры
CatBoostClassifier_f1 = f1_score(y_test, y_pred, average='macro')
print("CatBoostClassifier_f1: ", CatBoostClassifier_f1)
```

Output:
```
CatBoostClassifier_best_params:  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}
CatBoostClassifier_f1:  0.7447331848344864
```

<!-- grade_id: cell-bc10a3a8366a4951, points: 2, grade: true, solution: false -->
```python
assert type(model) == cb.CatBoostClassifier
assert type(grid_search) == GridSearchCV
```

```python
# Вывод результатов моделей классификации 
data = [ 
    ['LogisticRegression', f'{LogisticRegression_f1}', f'{LogisticRegression_best_params}'],
    ['SVC', f'{SVC_f1}', f'{SVC_best_params}'],
    ['DecisionTreeClassifier', f'{DecisionTreeClassifier_f1}', f'{DecisionTreeClassifier_best_params}'],
    ['RandomForestClassifier', f'{RandomForestClassifier_f1}', f'{RandomForestClassifier_best_params}'],
    ['GradientBoostingClassifier', f'{GradientBoostingClassifier_f1}', f'{GradientBoostingClassifier_best_params}'],
    ['XGBClassifier', f'{XGBClassifier_f1}', f'{XGBClassifier_best_params}'],
    ['CatBoostClassifier', f'{CatBoostClassifier_f1}', f'{CatBoostClassifier_best_params}']
]


# Создание DataFrame из данных
df = pd.DataFrame(data, columns=['Модель', 'F1', 'Параметры'], index=range(1, 1 + len(data)))
pd.set_option('display.max_colwidth', None)
# Отображение DataFrame
df.head(len(data))
```

Output:
```
Модель                  F1  \
1          LogisticRegression  0.6962137604356302   
2                         SVC   0.615523472539518   
3      DecisionTreeClassifier  0.7150713736547856   
4      RandomForestClassifier  0.7290506784614048   
5  GradientBoostingClassifier  0.7511473381048341   
6               XGBClassifier  0.7482487047186519   
7          CatBoostClassifier  0.7447331848344864   

                                                     Параметры  
1                                 {'C': 10.0, 'penalty': 'l2'}  
2                               {'C': 1.0, 'kernel': 'linear'}  
3                       {'criterion': 'gini', 'max_depth': 10}  
4                     {'max_depth': None, 'n_estimators': 200}  
5  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}  
6  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}  
7  {'learning_rate': 0.1, 'max_depth': 5, 'n_estimators': 100}
```